---
title: Sorting and grouping
description: Description

--- type:VideoExercise lang:r xp:50 skills:1 key:202c65c2f7
## Sort with ORDER BY

* ORDER BY
  * Single column
  * Multiple columns
  * DESC

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:NormalExercise lang:r xp:100 skills:1 key:8ee3754ce0
## Sort by a single column

Use ORDER BY on one column at a time, with and without DESC

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:b188642852
## Sort by multiple columns

Use  ORDER BY on multiple columns. Make tiebreaker analogy. Embed MCE to interpret single value from last result

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:5257ccf6a2
## Get started with GROUP BY

* GROUP BY
  * Single column
  * Multiple columns
* Combine with ORDER BY
* Introduce column aliases via AS as a convenient shorthand. Will become even more useful later in chapter

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:d50d14646f
## Everything in its right place

Q: Which of the following sorts column1 in ascending order within each level of column2?
A:
SELECT column1 AS c1, column2 AS c2
FROM table_name
GROUP BY c1
ORDER BY c2;

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:55599e8acd
## Sort within groups

Combine GROUP BY and ORDER BY to sort within groups. Embed another MCE at end to interpret last result

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:c87be43c26
## Combine GROUP BY with aggregate functions

* Unleash the full power of GROUP BY and aggregate functions by combining them to perform operations within groups
* Examples covering grouped COUNT, SUM, AVG

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:NormalExercise lang:r xp:100 skills:1 key:909839d3c8
## Compute group extremes

Combine GROUP BY with MIN and MAX to compute the extreme values within each group for a couple of variables

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:1bddabdab6
## Sort by group averages

Compute group averages, then sort by the averages

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:bdb1ebe87d
## Filtering with aggregate functions

* WHERE keyword cannot be used with aggregate functions
* HAVING
* Walk through several examples (this is not easy)

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:NormalExercise lang:r xp:100 skills:1 key:2e3d05e947
## HAVING patience

Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:a98b86006b
## Combine HAVING with ORDER BY

Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
