---
title: 'Case Study: Citi Bike In New York City'
description: >-
  This chapter gives you an opportunity to apply everything you've learned in
  the first three chapters to real-world data! You'll be working with a database
  that contains information on millions of individual bike share trips in New
  York City in 2015, along with the stations the bikes belong to, and daily
  weather metrics recorded from Central Park during the same time period.

--- type:VideoExercise lang:sql xp:50 skills:1 key:bc0d0e6e8d
## Introduction to the CitiBike Data

- Trips, Stations, Weather
- Imagine you're part of an urban planning team trying to understand biking habits in NYC to improve bike routes
- Etc.

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:sql xp:50 skills:1 key:8e3524960e
## Get to Know Your Data
Question goes here.

*** =instructions
-
-
-
-

*** =hint

*** =pre_exercise_code
```{sql}

```

*** =sct
```{sql}

```

--- type:TabExercise lang:sql xp:100 skills:1 key:feae971647
## Know Your LIMITs
- Selecting from the different tables, but using LIMIT
- Counts

*** =pre_exercise_code
```
Pre-exercise code goes here.
```
*** =sample_code
```
Sample code goes here.
```

*** =type1: NormalExercise
*** =instructions1
Instructions for first tab.
*** =solution1
```
Solution code goes here.
```
*** =sct1: SCT for first tab.

--- type:TabExercise lang:sql xp:100 skills:1 key:5f834c6a08
## Make Your Way Around (The Dataset)
Getting MIN, MAX, AVG

*** =pre_exercise_code
```
Pre-exercise code goes here.
```
*** =sample_code
```
Sample code goes here.
```

*** =type1: NormalExercise
*** =instructions1
Instructions for first tab.
*** =solution1
```
Solution code goes here.
```
*** =sct1: SCT for first tab.

--- type:TabExercise lang:sql xp:100 skills:1 key:48b9934c8f
## Filtering
Filtering based on various criteria

*** =pre_exercise_code
```
Pre-exercise code goes here.
```
*** =sample_code
```
Sample code goes here.
```

*** =type1: NormalExercise
*** =instructions1
Instructions for first tab.
*** =solution1
```
Solution code goes here.
```
*** =sct1: SCT for first tab.

--- type:TabExercise lang:sql xp:100 skills:1 key:c85b53bc9c
## More Filtering
More examples of filtering for interesting information

*** =pre_exercise_code
```
Pre-exercise code goes here.
```
*** =sample_code
```
Sample code goes here.
```

*** =type1: NormalExercise
*** =instructions1
Instructions for first tab.
*** =solution1
```
Solution code goes here.
```
*** =sct1: SCT for first tab.

--- type:TabExercise lang:sql xp:100 skills:1 key:aadfa696e6
## Weather
Some weather related questions

*** =pre_exercise_code
```
Pre-exercise code goes here.
```
*** =sample_code
```
Sample code goes here.
```

*** =type1: NormalExercise
*** =instructions1
Instructions for first tab.
*** =solution1
```
Solution code goes here.
```
*** =sct1: SCT for first tab.

--- type:TabExercise lang:sql xp:100 skills:1 key:f222b67a81
## Stations
Some questions related to stations

*** =pre_exercise_code
```
Pre-exercise code goes here.
```
*** =sample_code
```
Sample code goes here.
```

*** =type1: NormalExercise
*** =instructions1
Instructions for first tab.
*** =solution1
```
Solution code goes here.
```
*** =sct1: SCT for first tab.

--- type:TabExercise lang:sql xp:100 skills:1 key:6bda1c10dc
## Times
Some time related questions - e.g. most popular days to go cycling, etc.

*** =pre_exercise_code
```
Pre-exercise code goes here.
```
*** =sample_code
```
Sample code goes here.
```

*** =type1: NormalExercise
*** =instructions1
Instructions for first tab.
*** =solution1
```
Solution code goes here.
```
*** =sct1: SCT for first tab.

--- type:VideoExercise lang:sql xp:50 skills:1 key:6945e9c035
## Summary
- Summarize the case study
- Summarize the course
- Give suggestions for further projects using what you've learned

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8
