---
title: WHERE
description: >-
  With the WHERE statement, you can become very specific about which
  observations you want to pull out of your database. SELECT provides
  specificity to the variables or columns, WHERE provides the same thing for the
  observations.

--- type:VideoExercise lang:r xp: skills: key:a5b560fd80
## Selecting observations and combining predicates 
=, >, <, ... - AND, OR, NOT
Single quotes for text values, no quotes for numerics


*** =video_link

--- type:NormalExercise lang:r xp: skills: key:5cf99ebac5
## Equality 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp: skills: key:303e0f1293
## Advanced predicates 
BETWEEN, LIKE, IN

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:51c5b9aa48
## Title 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp: skills: key:e3dc9afc9e
## SELECT + WHERE = MAGIC 
Combine the WHERE clause with SELECT to answer more advanced problems (count the number of observations on the train y, what's the average age on passengers that go to Paris).

*** =video_link
