---
title: Introduction
description: |
  Get to know relational database management systems. 

--- type:VideoExercise lang:r xp: skills: key:98159ef95a
## What is an RDBMS? 
Explain the concept of a relational database, about SQL, and about queries. Give very simple example of a query: get a column from a database.

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:66a15333af
## Your first SQL query 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:MultipleChoiceExercise lang:r xp: skills: key:0cef8219f8
## Title 
Which of the following is a relational database management system, or RDBMS? (SQL, PostgreSQL, OracleDB, Microsoft SQL server). SQL is the language, not the RDBMS!

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp: skills: key:21c9839980
## SELECT 
How to select a column, multiple columns, all columns, and how to use aliases.
select a from x
select a, b from x
select * from x
select a as new_a, b as new_b from x


*** =video_link

--- type:NormalExercise lang:r xp: skills: key:0ab9e08abe
## Select all columns 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:add6959e06
## Select multiple columns 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:4370ce8e66
## Using aliases 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp: skills: key:2d416c919e
## SELECT (2) 
COUNT, SUM, AVG, MIN, MAX, DISTINCT

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:a9f4b7b1ef
## COUNT 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:2a3e02c440
## SUM and AVG 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:5fa5c4bc4e
## MIN and MAX 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:97b5f810ab
## DISTINCT 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
