---
title: WHERE
description: >-
  With the WHERE statement, you can become very specific about which
  observations you want to pull out of your database. SELECT provides
  specificity to the variables or columns, WHERE provides the same thing for the
  observations.

--- type:VideoExercise lang:r xp: skills: key:0cdce2a5ae
## WHERE predicates 
=, >, <, ... - AND, OR, NOT
Single quotes for text values, no quotes for numerics


*** =video_link

--- type:NormalExercise lang:r xp: skills: key:792ea9aba2
## Basic predicates (1) 
Use =, >, <

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:3c34c93fce
## Basic predicates (2) 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:MultipleChoiceExercise lang:r xp: skills: key:df5df94344
## Combining predicates (1) 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:3a9cc4c2eb
## Combining predicates (2) 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp: skills: key:e5d4ba36f1
## Advanced predicates 
BETWEEN, LIKE, IN

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:3bb8c05599
## BETWEEN 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:03b8a632ce
## IN 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:d51c29c111
## LIKE (1) 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:c5d1d4768a
## LIKE (2) 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp: skills: key:4a4190b1d6
## Combine SELECT and WHERE 
Combine the WHERE clause with SELECT to answer more advanced problems (count the number of observations on the train y, what's the average age on passengers that go to Paris).

*** =video_link

--- type:MultipleChoiceExercise lang:r xp: skills: key:c1098494e2
## Aggregate with basic predicate 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:ca79e1a93d
## Aggregate with combined predicate 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:c758e4fedd
## Aggregate with advanced predicate 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
