---
title: Introduction
description: |
  Get to know relational database management systems. 

--- type:VideoExercise lang:r xp: skills: key:5079d3acfd
## What is an RDBMS? 
Explain the concept of a relational database, about SQL, and about queries. Give very simple example of a query: get a column from a database.

*** =video_link

--- type:MultipleChoiceExercise lang:r xp: skills: key:f0e0c4c273
## Title 
Which of the following is a relational database management system, or RDBMS? (SQL, PostgreSQL, OracleDB, Microsoft SQL server). SQL is the language, not the RDBMS!

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:da972c11a6
## Your first SQL query 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp: skills: key:ef83760968
## SELECT 
How to select a column, multiple columns, all columns, and how to use aliases.
select a from x
select a, b from x
select * from x
select a as new_a, b as new_b from x


*** =video_link

--- type:NormalExercise lang:r xp: skills: key:37d2eb7ac3
## Select all columns 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:fa0052bb22
## Select multiple columns 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:d4ddf62caf
## Using aliases 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp: skills: key:e2a30621ec
## SELECT (2) 
COUNT, SUM, AVG, MIN, MAX, DISTINCT

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:c49ee243a0
## COUNT 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:2dc84ae2df
## SUM and AVG 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:b4a80da17c
## MIN and MAX 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:9f8d87c6d1
## DISTINCT 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
