---
title: WHERE
description: >-
  With the WHERE statement, you can become very specific about which
  observations you want to pull out of your database. SELECT provides
  specificity to the variables or columns, WHERE provides the same thing for the
  observations.

--- type:VideoExercise lang:r xp: skills: key:0e38bf0ea7
## WHERE predicates 
=, >, <, ... - AND, OR, NOT
Single quotes for text values, no quotes for numerics


*** =video_link

--- type:NormalExercise lang:r xp: skills: key:7ac581ac9c
## Basic predicates (1) 
Use =, >, <

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:9c35b41710
## Basic predicates (2) 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp: skills: key:5e39240624
## Advanced predicates 
BETWEEN, LIKE, IN

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:ff4a9cef3a
## Title 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp: skills: key:416d224297
## SELECT + WHERE = MAGIC 
Combine the WHERE clause with SELECT to answer more advanced problems (count the number of observations on the train y, what's the average age on passengers that go to Paris).

*** =video_link
