---
title: Filtering Rows
description: >-
  This chapter builds on the first by teaching you not only how to select and
  summarize columns of interest, but to filter tables for records satisfying
  some criteria of interest. You'll learn how to use basic comparison operators,
  combine multiple criteria, match patterns in text, and much more.

--- type:VideoExercise lang:sql xp:50 skills:1 key:4d7ac6f1f4
## WHERE oh WHERE
- Introduction to WHERE
- Filtering numeric values
- Filtering text values
- Basic comparison operators
- Multiple WHERE conditions: AND, OR

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:TabExercise lang:sql xp:100 skills:1 key:b90db25f33
## Simple Filtering of Text
Some basic filtering examples on text

*** =pre_exercise_code
```
Pre-exercise code goes here.
```
*** =sample_code
```
Sample code goes here.
```

*** =type1: NormalExercise
*** =instructions1
Instructions for first tab.
*** =solution1
```
Solution code goes here.
```
*** =sct1: SCT for first tab.

--- type:TabExercise lang:sql xp:100 skills:1 key:5bda32d7c8
## WHERE AND
Simple exercises combining WHERE with multiple AND clauses

*** =pre_exercise_code
```
Pre-exercise code goes here.
```
*** =sample_code
```
Sample code goes here.
```

*** =type1: NormalExercise
*** =instructions1
Instructions for first tab.
*** =solution1
```
Solution code goes here.
```
*** =sct1: SCT for first tab.

--- type:TabExercise lang:sql xp:100 skills:1 key:ecc1838fc7
## WHERE OR
Simple exercises combining WHERE, AND and OR clauses.

*** =pre_exercise_code
```
Pre-exercise code goes here.
```
*** =sample_code
```
Sample code goes here.
```

*** =type1: NormalExercise
*** =instructions1
Instructions for first tab.
*** =solution1
```
Solution code goes here.
```
*** =sct1: SCT for first tab.

--- type:VideoExercise lang:sql xp:50 skills:1 key:f13339152b
## More Advanced Filtering
Introduction to BETWEEN, IN - IS NULL, IS NOT NULL

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:TabExercise lang:sql xp:100 skills:1 key:9c11f67712
## BETWEEN (A Rock and A Hard Place)
Simple BETWEEN exercises

*** =pre_exercise_code
```
Pre-exercise code goes here.
```
*** =sample_code
```
Sample code goes here.
```

*** =type1: NormalExercise
*** =instructions1
Instructions for first tab.
*** =solution1
```
Solution code goes here.
```
*** =sct1: SCT for first tab.

--- type:TabExercise lang:sql xp:100 skills:1 key:4fc7e638f8
## WHERE IN The World
Exercises using WHERE with IN

*** =pre_exercise_code
```
Pre-exercise code goes here.
```
*** =sample_code
```
Sample code goes here.
```

*** =type1: NormalExercise
*** =instructions1
Instructions for first tab.
*** =solution1
```
Solution code goes here.
```
*** =sct1: SCT for first tab.

--- type:TabExercise lang:sql xp:100 skills:1 key:84411d78aa
## BETWEEN and AND
Exercises using BETWEEN with AND

*** =pre_exercise_code
```
Pre-exercise code goes here.
```
*** =sample_code
```
Sample code goes here.
```

*** =type1: NormalExercise
*** =instructions1
Instructions for first tab.
*** =solution1
```
Solution code goes here.
```
*** =sct1: SCT for first tab.
