---
title: GROUP BY / ORDER BY
description: Description

--- type:VideoExercise lang:r xp: skills: key:905b96b360
## GROUP BY 
'single' groups

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:b85525bc26
## Grouped Count (1) 
COUNT with GROUP BY

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:f393ba4e23
## Grouped Count (2) 
COUNT with GROUP BY with more advanced thing to group by GROUP BY

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp: skills: key:90836df839
## GROUP BY (2) 
'double' groups: two GROUP BY's and corresponding aggregators. 

*** =video_link

--- type:MultipleChoiceExercise lang:r xp: skills: key:adb9af5f21
## Advanced GROUP BY (1) 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:d082ea3af6
## Advanced GROUP BY (2) 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:43b29d0f94
## Advanced GROUP BY (3) 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp: skills: key:5c4925979b
## ORDER BY 
Talk about single order by, reversed order by, and two variables to order by (second one is the 'tie breaker')

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:1c67905bd6
## ORDER BY 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:aca96642d2
## Change the order 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:5b641e5416
## ORDER BY multiple columns 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
