---
title: Selecting and summarizing columns
description: This is a template chapter.

--- type:NormalExercise lang:r xp:100 skills:1 key:ac86c7bab9
## Template exercise

This is a template exercise.

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
