---
title: GROUP BY / ORDER BY
description: Description

--- type:VideoExercise lang:r xp: skills: key:e426d5d4bc
## GROUP BY 
'single' groups

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:2ad0822d93
## Grouped Count (1) 
COUNT with GROUP BY

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:b1a3b4547a
## Grouped Count (2) 
COUNT with GROUP BY with more advanced thing to group by GROUP BY

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp: skills: key:c47a01ae21
## GROUP BY (2) 
'double' groups: two GROUP BY's and corresponding aggregators. 

*** =video_link

--- type:MultipleChoiceExercise lang:r xp: skills: key:302ce24038
## Advanced GROUP BY (1) 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:da61dbb078
## Advanced GROUP BY (2) 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:7f7efa8109
## Advanced GROUP BY (3) 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp: skills: key:2705a84c7f
## ORDER BY 
Talk about single order by, reversed order by, and two variables to order by (second one is the 'tie breaker')

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:abece11868
## ORDER BY 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:1996cbd4a3
## Change the order 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:1dee073f37
## ORDER BY multiple columns 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
